---
title: "R Interactive Tutorial"
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
gradethis::gradethis_setup()
library(tidyverse)
library(gapminder)
knitr::opts_chunk$set(echo = FALSE)
```

## Ice Breaker

### Jeopardy
*NOTE: If you're taking this tutorial on your own you're welcome to press ***Continue*** *below to skip this warm-up.*

Before we get started let's play a game to warm up for the day! A game of Jeopardy. Don't worry about the reverse question answering style of the game. Answers are just fine! 

```{r embed, out.width = '100%'}
knitr::include_url("https://jeopardylabs.com/play/movies-tv-shows-42", height = "650px")
```

### #Goals
My goals is that by the end of this webinar tutorial you will be able to quickly generate information-packed graphs like the one below with only 2 lines of code.

```{r goals, exercise = TRUE}
gapminder_2007 <- filter(gapminder,year==2007)
ggplot(gapminder_2007,mapping = aes(x=gdpPercap,y=lifeExp,color = continent,size = pop)) + geom_point(alpha = 0.7)
```

## The Basic ingRedients

### Getting Started
For this tutorial we are going to imagine we are all new chefs in a fancy kitchen and we will be making a special dinner. To start "cooking" up code in R you'll need to know basic skills and ingredients you might find in a kitchen.

Write the R code required to add **two plus two**.   
When you want to see what the code does hit the **Run** button.  
If you get stuck hit the **Hint** button. Sometimes there may be multiple **Hints**.  
Finally, when you're ready to submit your answer press **Submit Answer**.


```{r two-plus-two, exercise=TRUE}

```

```{r two-plus-two-hint-1}
2
```

```{r two-plus-two-hint-2}
2 +
```

```{r two-plus-two-solution}
2 + 2
```

```{r two-plus-two-code-check}
grade_code()
```

### Subtraction

See how R performs other basic arithmetic. Try subtraction now. Perform **two minus two**.

**Remember** When you want to see what the code does hit the **Run** button.  
If you get stuck hit the **Hint** button. Sometimes there may be multiple **Hints**.  
Finally, when you're ready to submit your answer press **Submit Answer**.


```{r arithmetic, exercise=TRUE}


```

```{r arithmetic-hint-1}
2
```

```{r arithmetic-hint-2}
2 -
```

```{r arithmetic-solution}
2 - 2
```

```{r arithmetic-code-check}
grade_code()
```

### Multiplication

How do you think we would do multiplication? Try to **multiply 5 by 5**.

```{r multiply, exercise=TRUE}
#Write the multiplication expression below

```

```{r multiply-hint-1}
5
```

```{r multiply-hint-2}
5 *
```

```{r multiply-solution}
5 * 5
```

```{r multiply-code-check}
grade_code()
```

### More operations

Here are a couple more expressions you can use in R. Try to run the code!

```{r math_more, exercise=TRUE}
 # Division expression
5 / 2 

```

### More operations

And this last operation handles exponentiation. Try to run the code!

```{r math_more_2, exercise=TRUE}
# Exponentiation expression
2^5

```


### Making comments

See how comments affect the code. Using the "#" symbol at the beginning of a line prevents the code from being Run. In other words, R does not process this code. Click **Run** to see what happens when we put a # before the division operation below.

```{r comments, exercise=TRUE}
#Run the division operation below with the "#" in front
# 5 / 2 
```


### Test Your Knowledge

```{r quiz_comments, echo = FALSE}
question("What will the following code return?       
`#(5+4)/3*4-2`",
         answer("-1.25"),
         answer("10", message = "You remembered PEMDAS, but missed the most important symbol!"),
         answer("1.5"),
         answer("Nothing will be returned", correct = TRUE, message = "the # symbol means that code on this line will not be run"),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Why comments

So, we understand that "#" means the code won't be calculated. While the operation is not performed, we are effectively making a "comment" for ourselves. Commenting is how you take notes and give yourself a cue as to what the code's purpose is. Click **Run** to confirm that the answer we chose in the previous question was correct.

```{r why, exercise=TRUE, exercise.eval=TRUE}
#Run the code below. Adding a "#" at the beginning of a line prevents code from being run.

#(5+4)/3*4-2

```


### Variables

Now you will learn how to do variable assignment with the "<-" symbols. Assigning a number to a variable means this is a "numeric variable".

```{r variable, exercise=TRUE, exercise.lines = 7}
# Assign the value 37 to x
x <-  

```

```{r variable-solution}
x <- 37
```

```{r variable-code-check}
grade_code()
```

### Adding variables

Let's try adding variables together! What if I wanted to count the number of cats and dogs that I have at home? Set cats to two and dogs to three.

```{r variables_2, exercise=TRUE, exercise.lines = 12}
# Assign the values above to the variables cats and dogs
cats <- 
dogs <- 

# Add these two variables together
pets <-

# Below, type out the new variable you just created.

```

```{r variables_2-hint-1}
# Assign a value to the variables cats and dogs
cats <- 2
dogs <- 3

# Add these two variables together
pets <- 

# Below, type out the new variable you just created.

```

```{r variables_2-hint-2}
# Assign a value to the variables cats and dogs
cats <- 2
dogs <- 3

# Add these two variables together
pets <- cats + dogs

# Below, type out the new variable you just created.

```

```{r variables_2-solution}
# Assign a value to the variables cats and dogs
cats <- 2
dogs <- 3

# Add these two variables together
pets <- cats + dogs

# Below, type out the new variable you just created.
pets
```

```{r variables_2-code-check}
grade_code()
```

### Text variables

What else is a variable? Assigning a text string to a variable makes the variable a "character variable". Assign the phrase in quotation marks to the variable `text`, then underneath you need to type out `text` to display the value stored there. Once complete **Run** and then **Submit** your code.

```{r variables_3, exercise=TRUE, exercise.lines = 7}
# Assign a "string" of words to a variable
text <- "I have two cats one named Meeju and the other Daffne" 

# Type "text" below so R displays the variable value you just stored

```

```{r variables_3-solution}
# Assign a "string" of words to a variable
text <- "I have two cats one named Meeju and the other Daffne" 

# Type "text" below so R displays the variable value you just stored
text
```

```{r variables_3-code-check}
grade_code()
```

### Nomencature

Let's do some naming. I will start referring to the code block where you type out your code as the **script***. When you select **Run** you feed the script to the **console**.  

```{r nomenclature, exercise=TRUE}
# This code chunk here is called a "script".
# When you select Run, you feed the script to the "console"
"The console reads the script and gives you results"

```

### Test Your Knowledge

```{r quiz_variables, echo = FALSE}
question("What will the following code return?       
`cats <- 2
  dogs <- 5
  text <- 'I have two cats named Meeju and Daffne'
  pets <- cats + text
  pets`",
         answer("7"),
         answer("2, I have two cats named Meeju and Daffne"),
         answer("2 I have two cats named Meeju and Daffne"),
         answer("Error, you cannot add a number variable to character variable", correct = TRUE),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

## R-ecipes: Vectors and Functions:

### Vectors!
Now that we know our way around the kitchen and a couple ingredients, let's learn some basic recipes and some methods of cooking in R.

"Vectors" are a type of variable that holds a series of values. Like a shopping list or a receipt. You will be using a new "function" that looks like this `c()`. Values are separated by commas in the `c()` function. Assign any series of value to the variable - you can pick any numbers you want. Make sure to put a comma between the numbers though!

```{r vectors, exercise=TRUE, exercise.lines = 7}
#Assign a series of values to the my_receipt variable. 
my_receipt <- c()

#Tell the console to display values stored in the new variable your created by typing the variable name below

```

```{r vectors-hint-1}
#Assign a series of values to the my_receipt variable
my_receipt <- c(14,2,5,5,35,20,10,1)

#Tell the console to display values stored in the new variable your created by typing the variable name below

```

```{r vectors-solution}
#Assign a series of values to the my_receipt variable
my_receipt <- c(14,2,5,5,35,20,10,1)

#Tell the console to display values stored in the new variable your created by typing the variable name below
my_receipt
```

### More Vectors

Vectors can also be a series of texts values not just numbers. Remember you need to use " " with text. Below, a series of fruits are assigned to the `my_list` variable. All of you to do is type out the name of the variable where applicable. 

```{r vectors_2, exercise=TRUE, exercise.lines = 7}
#Assign a series of texts to the my_list variable, which will also be a vector
my_list <- c("apples", "bananas","eggplant","tofu", "pasta")

#Type out the name of the variable to see your values

```

### Vector Selection

You can use [] to select which element or value you want from each vector. **Run** the code below to see how this works.

```{r vectors_3, exercise=TRUE, exercise.lines = 7}
#Assign a series of texts to the my_list variable, which will also be a vector
my_list <- c("apples", "bananas","eggplant","tofu", "pasta")

#Now select the third element or object from the my_list vector
my_list[3]
```

### Test Your Knowledge

```{r quiz_brackets, echo = FALSE}
question("What will the result be from the following code?       
`my_receipt <- c(14,2,5,5,35,20,10,1)
my_list <- c('apples', 'bananas','eggplant','tofu', 'pasta')
my_receipt[1]`",
         answer("apples"),
         answer("Error, you cannot retrieve one element of numeric vector"),
         answer("5"),
         answer("14", correct = TRUE),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Vector functions
You can use apply functions to your vectors and variables. There is another function `mean()` that calculates the average value of a vector. All you have to do is "wrap" your variable name in the `mean()` function as you see below. Press **Run** to complete the calculation.
```{r mean, exercise=TRUE, exercise.lines = 7}
#Let's make a vector of a grocery shopping list
my_vector <- c(14, 5, 2, 12, 24, 100, 3, 400, 14, 14, 2, 20)

# Wrap the variable name in mean() to find the average value.
mean(my_vector)

```

### More Functions
There are countless functions to call, so try some! Just type out the name of a function seen below and put the variable name inside the parentheses.
```{r mean_2, exercise=TRUE}
#Let's use another shopping list vector again
my_vector <- c(14, 5, 2, 12, 24, 100, 3, 400, 14, 14, 2, 20)

# Try out other functions like median(), min(), max(), sd(), sort(), sum(), unique().

```

```{r mean_2-hint-1}
#Let's use another shopping list vector again
my_vector <- c(14, 5, 2, 12, 24, 100, 3, 400, 14, 14, 2, 20)

# Try out other functions like median(), min(), max(), sd(), sort(), sum(), unique().
median(my_vector)
```

### Saving the Result
We can also save the result of the function to another variable. Try saving the result of the `mean()` function for `my_vector` to yet another variable named `new_val`. Once complete, **Run** and then **Submit** your code.
```{r mean_3, exercise=TRUE}
#Let's use another shopping list vector again
my_vector <- c(14, 5, 2, 12, 24, 100, 3, 400, 14, 14, 2, 20)

#Now save the mean of the vector to a new variable
new_val <- mean()
```

```{r mean_3-solution}
#Let's use another shopping list vector again
my_vector <- c(14, 5, 2, 12, 24, 100, 3, 400, 14, 14, 2, 20)

#Now save the mean of the vector to a new variable
new_val <- mean(my_vector)
```

```{r mean_3-code-check}
grade_code()
```

## Making Dinner: Data Tables

### Dataframes
You are doing so well! You've learned all the basics of how to cook. But now we need to actually gather the ingredients and build the recipe for the main course. The main recipe is our data table.

A **dataframe** is what we call our data tables in R, like an excel table. Below type out the variable name `starwars` to display the A Star Wars related dataset. `starwars` is preloaded into our session and accessible for our use. Don't forget to hit **Run**
```{r df, exercise=TRUE}
#Type out starwars to display the R dataframe known as starwars
starwars
```

### Understanding your Data

It's useful to investigate our dataframe to understand it better. There are a number of functions to help us out. Let's start with the head() function. Once complete, **Run** and then **Submit** your code. What do you think this function does?
```{r df_head, exercise=TRUE}
#Apply the head() function to the dataframe like you would any other function and variable.
head()
```

```{r df_head-solution}
#Apply the head() function to the dataframe like you would any other function and variable.
head(starwars)
```

```{r df_head-code-check}
grade_code()
```

### More Understanding
We can also use another important function, glimpse(), to give us a general idea of what is stored in our dataframe. Wrap the `starwars` variable in the `glimpse()` function. Try it out below! 

Note: Each column of a dataframe is interestingly sometimes referred to as a "variable" in R. Can you see how that might make sense?
```{r df_glimpse, exercise=TRUE}
#Apply the glimpse() function to the dataframe like you would any other variable.
glimpse()

```

### Summary Function
Now let's try the `summary()` function too. Wrap `starwars` in the `summary()` function. Once complete, **Run** and then **Submit** your code.
```{r df_summary, exercise=TRUE}
#Apply the summary() function to the dataframe.

```

```{r df_summary-solution}
#Apply the head() function to the dataframe like you would any other function and variable.
summary(starwars)
```

```{r df_summary-code-check}
grade_code()
```

### Test Your Knowledge
```{r star_dataframe}
glimpse(starwars)
```

```{r quiz_dataframe, echo = FALSE}
question("Name a variable found in the starwars dataframe",
         answer("glimpse()"),
         answer("mean()"),
         answer("character"),
         answer("height", correct = TRUE),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

## RefreshR

We covered a lot so far so let's take a step back to catch our breath and refresh our knowledge. 

### Scattergories
*NOTE: If you're taking this tutorial on your own you're welcome to press ***Continue*** *below to skip this warm-up.*

Before we review, let's play a game to warm up for the day! A game that me and my friends like to play is called Scattergories. It's simple and fun. 

1. Everyone get out a pencil and paper or just take notes on a word document.
2. Make a list from 1 to 6 on the paper
3. You will be given 6 categories after I press the start button 
4. You have 1 minute think of a word or phrase for each category that STARTS with the big letter shown BUT the catch is that you only get points if no one else thought of that word!
5. At the end we will score our cards. Whoever has the most points at the end wins!

[Click here to play](https://swellgarfo.com/scattergories/)

### Refresher

Let's review what we've learned so far.

### Question 1
```{r refresher_1, echo = FALSE}
question("What code is needed to multiple 123 x 2?",
         answer("`246`", message = "You're great at math! But this is not the answer"),
         answer("`#123 x 2`", message = "the # symbol means that code on this line will not be run"),
         answer("`123 x 2`", message = "x is not the multiplication symbol"),
         answer("`123 * 2`", correct = TRUE),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Question 2
```{r refresher_2, echo = FALSE}
question("What code would make a new variable that adds together the variables 'fruits' and 'veggies' as seen below?     
` fruits <- 4
  veggies <- 7`",
         answer("`food <- 4 + 7`", message = "You added the numbers but not the variables."),
         answer("`food -> fruits + veggies`", message ="Look at the symbol you used to do variable assignment."),
         answer("`food == fruits + veggies`", message = "Look at the symbol you used to do variable assignment."),
         answer("`food <- fruits + veggies`", correct = TRUE),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Question 3
```{r refresher_3, echo = FALSE}
question("What function allows us create a vector from a series of values?",     
         answer("`mean()`", message = "This function gives you the average."),
         answer("`<-`", message ="The arrow is used for variable assignment"),
         answer("`select()`", message = "select() is used to choose columns in a dataframe"),
         answer("`c()`", correct = TRUE),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Question 4
```{r refresher_4, echo = FALSE}
question("What code do you need to find the average of my_vector?     
`my_vector <- c(14, 5, 2, 12, 24, 100, 3, 400, 14, 14, 2, 20)`",
         answer("`mean() <- my_vector`", message = "my_vector goes inside the mean function"),
         answer("`mean(c(14, 5, 2, 12, 24, 100, 3, 400, 14, 14, 2, 20))`", message ="You didn't take the mean of my_vector - you took the mean of the values."),
         answer("`average(my_vector)`", message = "average() is not a function"),
         answer("`mean(my_vector)`", correct = TRUE),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Question 5
```{r gapminder_refresher}
glimpse(gapminder)
```

```{r refresher_5, echo = FALSE}
question("Name a variable found in the gapminder dataframe",
         answer("summary()", message = "summary() summarizes our dataframe but is not a variable."),
         answer("c()", message = "This function combines values to make a vector"),
         answer("height", message = "You're thinking of the starwars dataframe!"),
         answer("lifeExp", correct = TRUE),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

## Prep the Recipe: Manipulating Data

### Starwars Dataframe
We've gathered the ingredients, made sure the food looked good, and examined how it all fits together. Now we can actually start dicing, spicing, and sauteing. 

In other words let's try "manipulating" the data. Back to starwars! Type the name of the dataframe `starwars` below and hit **Run** to display it.
```{r starwars, exercise=TRUE}
#Access the built-in starwars dataset by typing starwars

```

```{r starwars-solution}
#Access the built-in starwars dataset by typing starwars
starwars
```

### Leveling Up

We are going to need new functions to help us manipulate our data. To do this you will need to download a package, in order to Upgrade R. Hit **Run** to process this code to install and deploy the **tidyverse**. 

Note: When you hit Run you'll see the following message `Warning: package 'tidyverse' is in use and will not be installed`. This is fine - **tidyverse** is already loaded in our session. I just want you to see how you install and load new packages if you need to do it on your own.

```{r starwars_package, exercise=TRUE}
#The code below will download the "tidyverse" package from the internet
install.packages("tidyverse")

#Now we need to "load" the package. Type tidyverse into the library function
library(tidyverse)

```

### New Functions

Now that R is leveled up we can use new functions like `select()`, `filter()`, `arrange()`, and `mutate()`. Let's try out `select()`.

### Selecting

`select()` takes the following format `select(dataframe,column name, column name, column name, etc)`.

### What do you think?
```{r star_dataframe_2}
glimpse(starwars)
```

```{r ques_select, echo = FALSE}
question("What is the code to select the variables name and height from starwars?",
         answer('`select(starwars,name,height)`', correct = TRUE),
         answer('`select(starwars,name)`'),
         answer('`select(dataframe,1,2)`'),
         answer('`select(name,height)`'),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Select Starwars

Great! Now we can select columns from the starwars dataset. Let's **select** a couple columns to make it interesting. Hit **Run** to apply the `select()` function as seen below.
```{r starwars_select, exercise=TRUE}
# Select name, height, homeworld, and species
select(starwars,name, height,homeworld,species)

```

### Nested Selecting

Can we  put a function within a function? Why yes! In the code below wrap the `select()` function with the `head()` function. In functional terms, use `head()` to find the first 6 rows of the selected starwars columns. Once complete, **Run** and then **Submit** your code.  

If you get stuck hit the **Solution** button.

```{r starwars_select_2, exercise=TRUE}
#Try to wrap this whole function in the head() function.
select(starwars,name, height,homeworld,species)
```

```{r starwars_select_2-solution}
#Try to wrap this whole function in the head() function.
head(select(starwars,name, height,homeworld,species))
```

```{r starwars_select_2-code-check}
grade_code()
``` 

### Filtering

Sometimes we only want very specific data so we will need to **filter** it, just like in an excel filter. 
`filter()` takes the folllowing format. 
`filter(dataset, column_name [filtering_condition])`  

Take note of the code below. See how we created a new variable `starwars_head` and fed that into the filter. In sum, this code script will filter the first 6 rows of the `starwars` so that only characters whose homeworld is Tatooine will be shown. *Notice how we have to use `==` instead of `=` when doing a filtering condition inside of `filter()`*.  

Press **Run** to see the result of the `filter()` function on `starwars_head`.

```{r starwars_head}
starwars_head <- head(select(starwars,name, height,homeworld,species))
```

```{r starwars_filter, exercise=TRUE, exercise.setup = "starwars_head"}
#Let's create a new variable to save the "head" of starwars
starwars_head <- head(select(starwars,name, height,homeworld,species))

#Now apply the filter verb so that we only look at characters whose homeworld is Tatooine
filter(starwars_head, homeworld == "Tatooine")
```

### Test your Understanding
```{r star_filter_setup}
starwars_head <- head(select(starwars,name, height,homeworld,species))
starwars_head
```

```{r ques_filter, echo = FALSE}
question("How can we filter the starwars_head data to only show data for Droids?",
         answer('`filter(starwars_head,species == "Droid")`', correct = TRUE),
         answer('`filter(starwars_head,species = "Droid")`'),
         answer('`filter(species == "Droid")`'),
         answer('`filter(species == Droid)`'),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Answer
Let's try out that solution now and compare it to starwars_head. Below you first save `starwars_head`, second you will name `starwars_head` so that it displays to the console, and then lastly you will apply the `filter()` as seen in the answer above to get the filtered dataset.
```{r starwars_answer, exercise=TRUE}
#Save starwars_head again below so you can compare it to the filtered dataset next
starwars_head <- head(select(starwars,name, height,homeworld,species))

#Type out starwars_head so that the saved value is now "printed" or displayed in the console
starwars_head

#Now apply the filter verb so that we only look at characters whose species is Droid
filter(starwars_head,species == "Droid")

```

### More Filtering
Filtering is very flexible. Let's see how we can filter height for characters below 100. I'm showing you `starwars_head` automatically so it is easier to figure out the correct code.  

In the script below you need to figure out how to `filter()` `starwars_head` so the dataframe is limited to characters whose height is less than 100. Don't forget you can use **Hints** if you get stuck.  

Once complete **Run** and then **Submit** your code.
```{r filter_logic_head}
#Display starwars head as we have done before
starwars_head <- head(select(starwars,name, height,homeworld,species))
starwars_head
```

```{r filter_logic, exercise=TRUE, exercise.setup = "starwars_head"}
#Save starwars_head again below so you can apply the filter()
starwars_head <- head(select(starwars,name, height,homeworld,species))

#Now apply the filter verb so that height is less than 100

```

```{r filter_logic-hint-1}
#Save starwars_head again below so you can apply the filter()
starwars_head <- head(select(starwars,name, height,homeworld,species))

#Now apply the filter verb so that height is less than 100
filter(starwars_head,)
```

```{r filter_logic-hint-2}
#Save starwars_head again below so you can apply the filter()
starwars_head <- head(select(starwars,name, height,homeworld,species))

#Now apply the filter verb so that height is less than 100
filter(starwars_head, height < ) 
```

```{r filter_logic-solution}
#Save starwars_head again below so you can apply the filter()
starwars_head <- head(select(starwars,name, height,homeworld,species))

#Now apply the filter verb so that height is less than 100
filter(starwars_head, height < 100)
```

```{r filter_logic-code-check}
grade_code()
```

### Test your Understanding

```{r star_doublefilter_setup}
starwars_head <- head(select(starwars,name, height,homeworld,species))
starwars_head
```


```{r ques_doublefilter, echo = FALSE}
question("How can we filter starwars_head to only show data for Humans that are above 200 height.",
         answer('`filter(starwars_head,species == "Human", height > 200)`', correct = TRUE),
         answer('`filter(starwars_head,species == "Human" AND height > 200)`'),
         answer('`Impossible. You cannot use two filters in one function`'),
         answer('`filter(starwars_head,species == Human & height > "200")`'),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Answer
Let's try out that solution.
```{r starwars_df_answer, exercise=TRUE, exercise.setup = "starwars_head"}
#Apply the two filters to starwars_head
filter(starwars_head,species == "Human", height > 200)
```

### Arrange
We can also use the function arrange to sort data within a column. 
Arrange take the following format.  
`arrange(dataset, column name)`  

I've pre-stored `starwars_head` below so you can just type `starwars_head` access this dataframe. Hit **Run** to see how arrange works.

```{r starwars_arrange, exercise=TRUE, exercise.setup = "starwars_head"}
#Apply the function arrange to the height column
arrange(starwars_head)
```
### Descending
The desc() function reverses the order of arrange. Try wrapping `height` in the `desc()` function so `arrange()` gets the cue to reverse the order.

Don't forget you can use **Hints** if you get stuck. Once complete **Run** and then **Submit** your code.
```{r starwars_desc, exercise=TRUE, exercise.setup = "starwars_head"}
#Apply the desc() function to our arrange code.
arrange()
```

```{r starwars_desc-hint-1}
#Apply the desc() function to our arrange code
arrange(starwars_head,height)
```

```{r starwars_desc-hint-2}
#Apply the desc() function to our arrange code
arrange(starwars_head,(height))
```

```{r starwars_desc-solution}
#Apply the desc() function to our arrange code
arrange(starwars_head,desc(height))
```

```{r starwars_desc-code-check}
grade_code()
```
### Mutate
Yet another critical function is mutate(). This function allows you to modify existing columns, calculate new columns, and more. 

Mutate takes the following format: `mutate(dataset,new or existing column name = some calculation)`

I have a suspicion that the height provided is given in cm. So I am going to add a new column to this dataframe -- height in inches! Examine the script below, then hit **Run** to see how the new dataframe compares against the original `starwars_head` dataframe.
```{r starwars_unmutate}
#Run starwars_head to compare against new mutated dataframe
starwars_head
```

```{r starwars_mutate, exercise=TRUE, exercise.setup = "starwars_head"}
#Use mutate() to calculate a new column called height_inch. Hint: there are 2.54 cm in an inch
mutate(starwars_head,height_inch = height / 2.54)

```
### More Mutations
We can also use mutate to modify existing columns rather than make new columns. How could we use mutate to change starwars_head so that height is simply expressed in inches in the "height" column without making a new column?

Remember, Mutate takes the following format: `mutate(dataset,new or existing column name = some calculation)` and that there are 2.54 cms in an inch. So, set `height` equal to the `some calculation` portion of your code.  

Don't forget you can use **Hints** if you get stuck. Once complete **Run** and then **Submit** your code.
```{r starwars_unmutate_2}
#Run starwars_head to compare against new mutated dataframe
starwars_head
```

```{r starwars_mutate_2, exercise=TRUE, exercise.setup = "starwars_head"}
#Use mutate() to re-calculate the height column.

```

```{r starwars_mutate_2-hint-1}
#Use mutate() to re-calculate the height column.
mutate()
```

```{r starwars_mutate_2-hint-2}
#Use mutate() to re-calculate the height column.
mutate(starwars_head,height = )
```

```{r starwars_mutate_2-solution}
#Use mutate() to re-calculate the height column.
mutate(starwars_head,height = height / 2.54)
```

```{r starwars_mutate_2-code-check}
grade_code()
```

### Aggregating Data
Briefly, we will explore how to summarize and aggregate our data to make it easier to intepret(and graph!). Now let's try to see how many of each species type there are by using the count() function.
Count takes the following format: `count(dataset,column name)`

Below, apply the `count()` function to the original `starwars` dataframe to count how many species types there are. Hit the **Run** button to process the script. How many species types do you count?

```{r starwars_count, exercise=TRUE}
#Apply the count function to the original starwars dataframe
count(starwars,species)
```

### More Counting
Count takes a special "argument" called "sort". An **argument** is code that modifies a function. So just place a comma after species and write sort=TRUE like so:  
`count(dataset,column name, sort = TRUE)`  

In our case, `sort` is an **argument** of the `count()` function. Below you will need to run `count()` again on species and then set the `sort` argument to TRUE. Don't forget you can use **Hints** if you get stuck. Once complete **Run** and then **Submit** your code. 
```{r starwars_count_sort, exercise=TRUE}
#Apply the count function again but with Sort = TRUE 

```

```{r starwars_count_sort-hint-1}
#Apply the count function again but with Sort = TRUE 
count(starwars,species)
```

```{r starwars_count_sort-hint-2}
#Apply the count function again but with Sort = TRUE 
count(starwars,species,sort = )
```

```{r starwars_count_sort-solution}
#Apply the count function again but with Sort = TRUE 
count(starwars,species,sort = TRUE)
```

```{r starwars_count_sort-code-check}
grade_code()
```

### Now you try
What code do we need to count how many starwars characters are from each homeworld and have the order sorted from most to least frequent?   

Don't forget you can use **Hints** if you get stuck. Once complete **Run** and then **Submit** your code. You'll need the correct answer to solve the question after this exercise.
```{r starwars_count_planet, exercise=TRUE}
#Count the number of homeworlds in starwars and sort it
```

```{r starwars_count_planet-hint-1}
#Count the number of homeworlds in starwars and sort it
count()
```

```{r starwars_count_planet-hint-2}
#Apply the count function again but with Sort = TRUE 
count(starwars,homeworld,sort = )
```

```{r starwars_count_planet-solution}
#Apply the count function again but with Sort = TRUE 
count(starwars,homeworld,sort = TRUE)
```

```{r starwars_count_planet-code-check}
grade_code()
```

### Planet Problem

```{r ques_count_homeworld, echo = FALSE}
question("Which homeworld are the most starwars characters from in the `starwars` dataframe",
         answer('Naboo', correct = TRUE),
         answer('Tatooine', message = "Not quite, 10 characters are from Tatooine. 11 characters are from this planet."),
         answer('48', message = "It's true that there are 48 unique homeworlds but you should reread the question"),
         answer('Alderaan', message = "You may have forgotten to sort your dataframe"),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Summarizing data
Another crucial aggegating function is summarize. In this exercise, we will attempt to find the minimum height of the starwars dataframe. 

summarize() takes the following format: `summarize(dataframe, function(existing column)`.

Use `summarize()` to find the **minimum height** of the starwars dataframe. Once complete **Run** and then **Submit** your code. The result **may surprise you**.
```{r starwars_summarize, exercise=TRUE}
#Use the summarize() and min() function to gives us the minimum height of the starwars characters

```

```{r starwars_summarize-hint-1}
#Use the summarize() and min() function to gives us the minimum height of the starwars characters
summarize()
```

```{r starwars_summarize-hint-2}
#Use the summarize() and min() function to gives us the minimum height of the starwars characters
summarize(starwars,min())
```

```{r starwars_summarize-solution}
#Use the summarize() and min() function to gives us the minimum height of the starwars characters
summarize(starwars,min(height))
```

```{r starwars_summarize-code-check}
grade_code()
```

### NA issue
Yep that's right, you were suppose to get "NA" as an answer. I suspect that we cannot find the minimum height of the height variable because one or multiple characters have an undefined height. In other words some characters have their height listed as "NA".

Let's double check that. We can either go through the entire starwars dataframe to look for NA as height or we can do a shortcut. **Run** the following script to see all entries for height. Observe how `$` is used access the height variable from starwars. Do you see any NAs?

```{r starwars_height_examine, exercise=TRUE}
#See all entries for height
starwars$height
```

### Dropping NA's
Yes there were 6 NAs, which is why the minimum could not be calculated. Let's see if we tell minimum to disregard NAs.  

The min() function has a special argument called `na.rm` which means **r**e**m**ove **NA**s. So this time let's try the `summarize()` function again to find the minimum height, but this time set `na.rm` to TRUE inside of min() like so:  
`min(height,na.rm=TRUE)` 

And remember `summarize()` takes the following format: `summarize(dataframe, function(existing column)`.
```{r starwars_dropna, exercise=TRUE}
#Now try the summarize() function again but this time remove the NAs from height.

```

```{r starwars_dropna-hint-1}
#Now try the summarize() function again but this time remove the NAs from height.
summarize()
```

```{r starwars_dropna-hint-2}
#Now try the summarize() function again but this time remove the NAs from height.
summarize(starwars,min(height))
```

```{r starwars_dropna-solution}
#Now try the summarize() function again but this time remove the NAs from height.
summarize(starwars,min(height,na.rm = TRUE))
```

```{r starwars_dropna-code-check}
grade_code()
```

### More Summarizing
Now that we know how to eliminate NA's let's find the **maximum height** of the starwars characters. **Run** the code below and then solve the question related to it. Use the `max()` function inside `summarize()` to find the *maximum height**
```{r starwars_summarizing, exercise=TRUE}
#Use max() to find the max height of starwars

```


```{r ques_max_height, echo = FALSE}
question("What is the maximum height for the starwars characters",
         answer('264', correct = TRUE),
         answer('66', message = "You found the minimum height. Find the max height now."),
         answer('NA', message = "You forgot to remove NAs."),
         answer('202', message = "It seems like you found the max height of starwars_head and not starwars"),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Starwars Trivia
And now just for fun, answer the following off the top of your head:  
The smallest starwars character is(66cm)?  
The tallest starwars character is(264cm)?

### Trivia Answers
Smallest - Yoda (66 cm) \
Tallest - Yarael Poof (264 cm)

![Yoda](images/yoda.png){#id .class width=50% height=50%}
![Yarael Poof](images/poof.png){#id .class width=50% height=50%}

### Renaming Columns
You can use the `rename()` function to rename columns in your dataframe. Rename takes the following format:  
`rename(dataframe,column new name = column old name)` \

Rename the height column to read "height_cm" rather than just "height".
```{r starwars_rename, exercise=TRUE}
#Apply the rename() function to starwars and change height to height_cm

```

```{r starwars_rename-hint-1}
#Apply the rename() function to starwars and change height to height_cm
rename(starwars)
```

```{r starwars_rename-hint-2}
#Apply the rename() function to starwars and change height to height_cm
rename(starwars,height_cm)
```

```{r starwars_rename-solution}
#Apply the rename() function to starwars and change height to height_cm
rename(starwars,height_cm = height)
```

```{r starwars_rename-code-check}
grade_code()
```

### Checkpoint: Special Code

**Congratulations!** You're all caught up with the rest of the class now.   

You have unlocked the special code: **Katherine Johnson**   

You can **stop** the tutorial and take the remainder of the course in-person. Please send an email with the subject line **Katherine Johnson** to me, Zach Margulies, by clicking [here](mailto:zachary.margulies@baltimorecity.gov?subject=Katherine%20Johnson). This will confirm to me that you are all caught up. You're welcome to continue if you wish, but we will cover the rest in person. 

## It's all about pResentation: Graphing

### Just the basics
You finished cooking and it smells amazing. The work is not done though, you still have to plate the food and make sure it looks presentable. 

R can quickly and easily produce publication quality graphs. There are a lot of ways to tweak graphs but let's just keep it simple for now.

In the script below we will apply the `head()` function to a dataset known as `gapminder` that I have preloaded for us. Click **Run** to process the script and inspect `gapminder`

```{r gapminder_1, exercise=TRUE}
#I've installed and loaded the gapminder dataset for us. Let's check it out with head()
head(gapminder)
```

### Investigate
Of note "gdpPercap" means Gross Domestic Product Per Capita. 

So, GDP per Capita is a metric for a **country's economic output per person("capita")** over one year. **GDP per Capita** is calculated by **dividing** the **GDP** of a country by its **population**. Some use this metric as an indicator of the economic health or status of a county. What other functions besides head() might be useful to explore this dataset? 

Let's investigate the dataset(hereafter referred to as a dataframe) further. Try applying the `summary()` function to the gapminder dataframe to better understand the data. Once complete **Run** and then **Submit** your code.

```{r gapminder_2, exercise=TRUE}
#Use the summary() function on the gapminder dataframe

```
```{r gapminder_2-hint-1}
#Use the summary() function on the gapminder dataframe
summary
```

```{r gapminder_2-hint-2}
#Use the summary() function on the gapminder dataframe
summary()
```

```{r gapminder_2-solution}
#Use the summary() function on the gapminder dataframe
summary(gapminder)
```

```{r gapminder_2-code-check}
grade_code()
```

### Relationships
Over the next few sections we will be investigating whether or not there is a  relationship(or correlation) between **life expectancy(lifeExp)** and **GDP per Capita(gdpPercap)**. To make things simpler, let's say I only want to investigate this for the year **2007**. Let's begin.

### Test out Understanding
Let's start by looking at data for the year 2007. Below I've loaded a `glimpse` of `gapminder` to help you with the question.

```{r ques_2007_help}
#A glimpse at gapminder to help with the question below
glimpse(gapminder)
```

```{r ques_2007, echo = FALSE}
question("How can we chop up the gapminder dataset to show data for the year 2007",
         answer('`filter(gapminder,year == 2007)`', correct = TRUE),
         answer('`filter(gapminder,year = 2007)`', message = "Remember the formatting of the filter symbols"),
         answer('`select(gapminder,year,2007)`', message ="You have not filtered year by anything and 2007 is not a variable in the dataframe."),
         answer('`count(starwars,year == 2007)`',message="Remember your dataframe name!"),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Filtering
Great! Let's see what the correct answer looks like below. Click **Run**. 
```{r gapminder_3b, exercise=TRUE}
filter(gapminder,year==2007)
```

### Summary
To help us explore the characteristics of the filtered dataframe let's try `summary()` again. 

We'll start by saving the filtered function above to a new variable, which we will name `gapminder_2007`. Go ahead and type this below.

Don't forget you can use **Hints** if you get stuck. Once complete **Run** and then **Submit** your code. 
```{r gapminder_2007, exercise=TRUE}
# Filter gapminder for the year 2007 and save it to a new variable, gapminder_2007

```

```{r gapminder_2007-hint-1}
# Save code above to a new variable, gapminder 2007
filter(gapminder,year==2007)
```

```{r gapminder_2007-hint-2}
# Save code above to a new variable, gapminder 2007
<- filter(gapminder,year==2007)
```

```{r gapminder_2007-solution}
# Save code above to a new variable, gapminder 2007
gapminder_2007 <- filter(gapminder,year==2007)
```

```{r gapminder_2007-code-check}
grade_code()
```

### Summary
Now that we have the filtered dataframe saved to a variable, just use the `summary()` function on this new variable. Click **Run** on the code below to see the results of this. Inspect the results.
```{r gapminder_summary,exercise = TRUE}
#Save code above to a new variable, gapminder 2007
gapminder_2007 <- filter(gapminder,year==2007)

#Use summary() on gapminder_2007
summary(gapminder_2007)
```

### Graphing! Finally!
Now we can finally begin to make graphs!! The way we do that is by using a highly customizable set of functions in the **ggplot2 package**.Thankfully ggplot2 is a part of the **tidyverse package**, which we already installed. Below I will show you basic code to setup a standard scatterplot with ggplot2. 

Click **Run** on the code below to see what the code does. Inspect the results. What do you notice and what do you think of this graph? What is good and what is bad?

```{r gapminder,echo=FALSE}
#Save code above to a new variable, gapminder 2007
gapminder_2007 <- filter(gapminder,year==2007)
```

```{r gapminder_5, exercise=TRUE, exercise.setup = "gapminder"}
#Run the following code
ggplot(gapminder_2007,mapping = aes(x=lifeExp,y=gdpPercap)) + 
  geom_point()
```

### Break it down
We made a scatterplot! I know it's a weird format but you'll catch on. Let's break it down piece by piece. 

First we need to run the `ggplot()` function on our filtered dataframe, `gapminder_2007`. Click **Run** on the code below to see what the code does. 
```{r gapminder_6, exercise=TRUE, exercise.setup = "gapminder"}
#Try Running the code without the additional parameters.
ggplot(gapminder_2007)
```

### Aesthetics
Running ggplot on the dataframe alone just produced an empty plot. Next we need to add a new argument to ggplot() which is `mapping = aes()`. Click **Run** on the code below.

Once again we have an empty plot.
```{r gapminder_7, exercise=TRUE, exercise.setup = "gapminder"}
#Try Running the code with the mapping = aes() argument
ggplot(gapminder_2007,mapping = aes())
```

### Don't forget your x and y's!
Whoops we forgot to add x and y values! The `aes()` function inside the `mapping` argument allows you to customize your ggplot. Here we can set what our x and y values ought to be. 

Before I said that I wanted to explore the potential relationship between Life Expectancy and GDP per capita. So let's set `x = lifeExp` and `y = gdpPercap` inside the `map = aes() argument`. Click **Run** on the code below. 

Now we have an empty plot with axis, but still no data points.
```{r gapminder_8, exercise=TRUE, exercise.setup = "gapminder"}
#Now try the code with the x and y values set.
ggplot(gapminder_2007,mapping = aes(x=lifeExp,y=gdpPercap))
```

### What's the point?
Something's not right here! We have the graph but still no points. Well we have to **add** the **geom_point()** function which tells ggplot to make a scatter plot. In other words, we need to place the food on plates!

In the code below you need to place a `+` after the `ggplot()` function and then type `geom_point()`. Click **Run** on the code below. 

Et Voilà now you have a scatterplot!
```{r gapminder_9, exercise=TRUE, exercise.setup = "gapminder"}
#Add the geom_point function to the ggplot() function
ggplot(gapminder_2007,mapping = aes(x=lifeExp,y=gdpPercap)) + 
  geom_point()
```

### Dependency 
Sometimes sticks out to me about this relationship. I feel like the x and y values could be different in some way. What do you think?

### Dependency 
The way I see it Life Expectancy is a dependent variable which depends on GDP per capita(the independent variable). So, `x=gdpPercap` and `y=lifeExp` makes more sense to me. In other words, I believe that a country's life expectancy is a function of the economic health of said country. In the code below, switch the x and y values so that `x=gdpPercap` and `y = lifeExp`.

Don't forget you can use **Hints** if you get stuck. Once complete **Run** and then **Submit** your code.

```{r gapminder_10, exercise=TRUE, exercise.setup = "gapminder"}
#Now try changing the x and y values.
ggplot(gapminder_2007,mapping = aes(x=lifeExp,y=gdpPercap)) + geom_point()
```

```{r gapminder_10-hint-1}
#Now try changing the x and y values.
ggplot(gapminder_2007,mapping = aes(x=gdpPercap,y=)) + geom_point()
```

```{r gapminder_10-solution}
#Now try changing the x and y values.
ggplot(gapminder_2007,mapping = aes(x=gdpPercap,y=lifeExp)) + geom_point()
```

```{r gapminder_10-code-check}
grade_code()
```

### Trends
Let's keep customizing! We will apply a trendline to the graph which helps to visualize trends in the data. Below, I am adding the function `geom_smooth()` to add a trendline. Click **Run** and inspect the results. 
```{r gapminder_11, exercise=TRUE, exercise.setup = "gapminder"}
#Run the following code
ggplot(gapminder_2007,mapping = aes(x=gdpPercap,y=lifeExp)) + 
  geom_point() + 
  #Add the geom_smooth() function to add a trend line to our graph
  geom_smooth()
```

### Adding Color
And here's another great way to customize our graph, by adding color. `aes()` is short for aesthetics so let's put `color` in that function call within the `ggplot()` function. Let's set `color = continent` so that color is associated with the continent variable. 

In the script below, set color equal to continent. 

You can click **Solution** if you get stuck. Once complete **Run** and then **Submit** your code. Inspect the results.

```{r gapminder_12b}
glimpse(gapminder_2007)
```

```{r gapminder_12, exercise=TRUE, exercise.setup = "gapminder"}
#Set color equal to continent
ggplot(gapminder_2007,mapping = aes(x=gdpPercap,y=lifeExp,color = )) + geom_point()
```

```{r gapminder_12-solution}
#Run the following code
ggplot(gapminder_2007,mapping = aes(x=gdpPercap,y=lifeExp,color = continent )) + geom_point()
```

```{r gapminder_12-code-check}
grade_code()
```

### More Aesthetics
Let's take a moment to consider other variables within this dataframe that we can **visualize**. Below we will are using `glimpse()` to look at `gapminder_2007` variables to help us think.

```{r gapminder_13a, include = TRUE}
glimpse(gapminder_2007)
```

### More Aesthetics
I think it makes sense to visualize the population size from`gapminder_2007`. One way to do that is by setting `size` is equal to `pop` inside of the `aes()` function after the `color` argument.

You can click **Solution** if you get stuck. Once complete **Run** and then **Submit** your code. Inspect the results.
```{r gapminder_13_b, exercise=TRUE, exercise.setup = "gapminder"}
#set size is equal to pop inside of aes()
ggplot(gapminder_2007,mapping = aes(x=gdpPercap,y=lifeExp,color = continent,size = )) + geom_point()
```

```{r gapminder_13_b-solution}
#set size is equal to pop inside of aes()
ggplot(gapminder_2007,mapping = aes(x=gdpPercap,y=lifeExp,color = continent, size = pop)) + geom_point()
```

```{r gapminder_13_b-code-check}
grade_code()
```

### Points everywhere
This is a lot information jammed into one graph! I am noticing a problem with some of my points. Take a moment to think about what might not be ideal about the plot and how you would change if you could magically alter the graph.
```{r gapminder_14, exercise=TRUE, exercise.setup = "gapminder",exercise.eval=TRUE}
#Run the following code
ggplot(gapminder_2007,mapping = aes(x=gdpPercap,y=lifeExp,color = continent,size = pop)) + geom_point()
```

### Points everywhere
The problem I am noticing is that many points overlap. ggplot2 has a way to fix that easily. Simply set the `alpha` argument to any value between 0 and 1. Make sure alpha is inside the `geom_point()` function since this is an argument specifically related to just the data points' characteristics. 

In the code below set `alpha` to any value between 0 and 1 to see how that affects the graph. **Click Run** to inspect the graph. 
```{r gapminder_14_b, exercise=TRUE, exercise.setup = "gapminder"}
#Set alpha equal to any number between 0 and 1
ggplot(gapminder_2007,mapping = aes(x=gdpPercap,y=lifeExp,color = continent,size = pop)) + 
  geom_point(alpha = )
```

### Progress
Amazing! Our exploratory graph is packed with information: life expectancy, GDP per capita, continents, and populations. 
```{r gapminder_15, exercise=TRUE, exercise.setup = "gapminder",exercise.eval = TRUE}
#Graph of gapminder_2007
ggplot(gapminder_2007,mapping = aes(x=gdpPercap,y=lifeExp,color = continent,size = pop)) + 
  geom_point(alpha = 0.7)
```

### A graph walks into a Bar...
There's so many types of graphs and we won't have time to cover them all. Let's try a bar graph. This time around, instead of using  `geom_point`, we will use `geom_col` which will produce bars(otherwise known as columns). 

In the code below you will need to set the `x = continent` and `y = pop`. 

Don't forget you can use **Hints** if you get stuck. Once complete **Run** and then **Submit** your code.
```{r bar_1, exercise=TRUE, exercise.setup = "gapminder"}
#Set x = continent and y = pop inside of aes()
ggplot(gapminder_2007,mapping = aes(x= ,y= )) + geom_col()
```

```{r bar_1-hint-1}
#Set x = continent and y = pop inside of aes()
ggplot(gapminder_2007,mapping = aes(x = continent,y= )) + geom_col()
```

```{r bar_1-solution}
#Set x = continent and y = pop inside of aes()
ggplot(gapminder_2007,mapping = aes(x = continent,y= pop)) + geom_col()
```

```{r bar_1-code-check}
grade_code()
```

### Scales
Inspect the graph we made. What do you notice about the scales? Is this easy to comprehend?

```{r bar_1_b}
#Set x = continent and y = pop inside of aes()
ggplot(gapminder_2007,mapping = aes(x = continent,y= pop)) + geom_col()
```

### Cheat code
I am going to give you an advanced cheat code to make the y value easier to interpret. Notice that I added a new line below `geom_col` that says `scale_y_continuous(labels = scales::comma)`. This *cheat code* converts the y-axis scale to a more understandable form. 

Click **Run** to see how the graph has changed.
```{r bar_2, exercise=TRUE, exercise.setup = "gapminder"}
#Set x = continent and y = pop inside of aes()
ggplot(gapminder_2007,mapping = aes(x = continent, y = pop)) + 
  #Here is the geom_col function again
  geom_col() +
  #You might not understand this next line, so you'll have to trust me.
  scale_y_continuous(labels = scales::comma)
```

### Pulling it all together
```{r ques_ggplot, echo = FALSE}
question("What is the best format to make a graph that compares lifeExp to year specifically for the country of Albania?",
         answer('`ggplot(filter(gapminder, country == "Albania"),mapping = aes(x=year,y=lifeExp)) + geom_point()`', correct = TRUE),
         answer('`ggplot(gapminder,mapping = aes(x=lifeExp,y=year)) + geom_point()`',message = "You didn't specify that the country is Albania"),
         answer('`ggplot(filter(gapminder, country = "Albania"),mapping = aes(x=year,y=lifeExp)) + geom_scatterplot()`', message = "geom_scatterplot is not a type of plot and theres another issue as well"),
         answer('`ggplot(filter(gapminder, country == "Albania",year ==2007),mapping = aes(x=year,y=lifeExp)) + geom_scatterplot()`', message = "Check the type of plot in your answer"),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

### Great Job
We had to use what we learned throughout this whole tutorial to nest functions, filter the gapminder dataframe, make a scatterplot, AND decide which variables should go where. Let's see the answer above graphed out now.  

`ggplot(filter(gapminder, country == "Albania"),mapping = aes(x=year,y=lifeExp)) + geom_point()`
```{r ques_ggplot_answer}
#Now run our answer 
ggplot(filter(gapminder, country == "Albania"),mapping = aes(x=year,y=lifeExp)) + 
  geom_point()
```

### Going Forward
That's a wrap! You completed the R tutorial - congratulations! But, before you go I want to show you how you can continue learning and practicing at home.

## Where in the world is R?
### Incredible Job!
You completed your training in the R-kitchen and served dinner. You might be asking yourself, how exactly can I cook at home like this? Keep reading!

### Is this R?
I know you might be wondering how in the world do I actually use R and where do I find it. Is it a website, is it a computer program, is what we are doing R? The answer is Yes. They are all R in some form. 

![](images/meme_butterfly.jpg){#id .class width=100% height=100%}

### RStudio Cloud quick tutorial
The easiest way for anyone to pick up and start using R is by going to [Rstudio Cloud](https://rstudio.cloud/) for a cloud based version of the R program. In the video below I will show you how to quickly start doing that. Thanks for completing this tutorial and good luck on your JouRney!
![Rstudio Cloud tutorial](https://youtu.be/zdiC4-ubqag)

### Expanding your knowledge
There are many ways to learn more cool stuff in R. I would recommend checking out [Datacamp's R courses](https://www.datacamp.com/courses/free-introduction-to-r). The first chapter of all their courses is free and then you can choose to pay $25/mo to unlock all of the datacamp courses. It's where I learned most of what I know. 

Alternatively, you can learn R for free by going to your [Rstudio Cloud](https://rstudio.cloud/) and typing in the console `install.packages("swirl")` and then `library(swirl)`.

[Rstudio Cloud](https://rstudio.cloud/) also has a lot of Guides and Primers to learn R. Check out the lefthand tabs on the Rstudio Cloud page for these guides and primers.

## Source

### Steal my code!
This tutorial was written by Zach Margulies[(zjmarg@gmail.com)](mailto:zjmarg@gmail.com). You can find the source code for this tutorial [here](https://github.com/zjmarg/Interactive-R-tutorial/blob/main/webinar_1/webinar_1.Rmd). This tutorial is intended to be open-source, freely available to the public, and user-friendly. So please download the code, modify it, and make your own tutorial.  

The inspiration for this tutorial came from many sources notably the brilliant teachers and coders at [Datacamp](https://learn.datacamp.com/) and the creators at the [learnr tutorials](https://rstudio.github.io/learnr/index.html). Special thanks to Gina, Alexandra, Jon, Joan, Sonney, Ricky, Amit, Dr. Cargill, Lin, and Nargis for your feedback, support, and kindness. 
